# Data Fundamentals (H)
John H. Williamson -- Session 2020/2021

----

**This submission must be your own work; you will have to make a Declaration of Originality on submission.**



```python
!unzip /content/week_1_intro.zip
```

Output:
```
Archive:  /content/week_1_intro.zip
   creating: week_1_intro/
   creating: week_1_intro/data/
  inflating: week_1_intro/data/cobra3_redux.obj  
 extracting: week_1_intro/data/parrots.png  
  inflating: week_1_intro/data/snails.txt  
  inflating: week_1_intro/data/sunspots.csv  
  inflating: week_1_intro/data/Wheat.csv  
   creating: week_1_intro/guides/
  inflating: week_1_intro/guides/test_array.txt  
   creating: week_1_intro/imgs/
  inflating: week_1_intro/imgs/1080ti.jpg  
 extracting: week_1_intro/imgs/castle.png  
 extracting: week_1_intro/imgs/df3.png  
  inflating: week_1_intro/imgs/Dungeon_0.gif  
  inflating: week_1_intro/imgs/faces_5.png  
  inflating: week_1_intro/imgs/impact.ttf  
  inflating: week_1_intro/imgs/join.png  
  inflating: week_1_intro/imgs/join.svg  
  inflating: week_1_intro/imgs/Low_poly.jpg  
  inflating: week_1_intro/imgs/notation.jpg  
 extracting: week_1_intro/imgs/orchid.png  
 extracting: week_1_intro/imgs/owl.png  
  inflating: week_1_intro/imgs/rank
```

---


# Lab 1
## Introduction to Juypter, Python and Numpy
### John H. Williamson 2022/2023
<div class="alert alert-block alert-success">


# Unassessed

</div>



This lab is intended to introduce you to the basic use of the Juypter+Python+Numpy environment, and show how the autograding in exercises works.You will get a mark (so you can see how autograding works), but this **will not count towards your final grade**. 

It is strongly recommended that you complete this exercise fully. This will take around two hours if you already know some NumPy.


## Purpose of this lab
This lab should help you:

* create simple arrays
* index and slice arrays
* stack arrays together
* compute simple statistics of arrays
* understand array arithmetic and broadcasting rules
* apply conditions to arrays and use Boolean indexing



# 1: Jupyter
If you have not used Jupyter before [jump to the Jupyter Quickstart](guides/JupyterGuide.ipynb) before continuing. If you have done Level 1 Computing Science, you do not need to do this.

## 2: Autograder tests

Lab exercises will (mainly) be autograded via automatic tests.

The following parts have some questions to answer, and some tests (which you cannot alter) which will be run against the code you have written. If the tests pass, you will see how many marks you got with a green tick. If they do not pass, you will see a red cross. Remember, this exercise doesn't count for anything, but do try to complete the exercises.

```python
# Make sure you run this cell!
from utils.tick import reset_marks, summarise_marks, marks
from utils.checkarr import array_hash, check_hash
import numpy as np  # NumPy
from utils.matrices import print_matrix, show_boxed_tensor_latex

# Set up Matplotlib
import matplotlib as mpl   
import matplotlib.pyplot as plt
%matplotlib inline

import utils.image_audio as ia

reset_marks()
print("Everything imported OK")

```

Output:
```
Everything imported OK
```

Here's a free 4 marks:

<!-- grade_id: cell-b12d0a22c191c138, points: 4, grade: true, solution: false -->
```python
with marks(4):
    print("Hello world")
```

Output:
```
Hello world
<IPython.core.display.HTML object>
```

And here's what happens when you have an error. Try setting `a` to 1, and making sure you can get this to pass. 

<!-- grade_id: cell-3040ccf7669e107d, points: 0, grade: false, solution: true -->
```python
a = 2
# YOUR CODE HERE
a = 1
```

<!-- grade_id: cell-221df43d91c7b63b, points: 3, grade: true, solution: false -->
```python
with marks(3):
    assert(a==1)
```

Output:
```
<IPython.core.display.HTML object>
```

----------------


# 4. Introduction to NumPy

We will be using [numpy](numpy.org) as the basis for our numerical operations. This provides a datatype called `ndarray`, that can be used to store and manipulate arrays of numbers.

<div class="alert alert-block alert-warning">
    
## NumPy worked example

If you have not used NumPy before or if you are rusty **[then work through the example before starting](guides/numpy_example.ipynb)**

</div>

---

## References and cheat sheets

If you are stuck, the following resources are very helpful:

### Cheatsheets
* [NumPy cheatsheet](https://github.com/juliangaal/python-cheat-sheet/tree/master/NumPy)
* [Python for Data Science cheatsheet](https://s3.amazonaws.com/assets.datacamp.com/blog_assets/PythonForDataScience.pdf)
* [Another NumPy Cheatsheet](https://s3.amazonaws.com/assets.datacamp.com/blog_assets/Numpy_Python_Cheat_Sheet.pdf)

### API reference and user guide
* [NumPy API reference](https://docs.scipy.org/doc/numpy-1.13.0/reference/)
* [NumPy user guide](https://docs.scipy.org/doc/numpy-1.13.0/user/basics.html)



-----
# NumPy

The foundation package for numerical operations is **NumPy** which provides an array type and accelerated operations on it. 

A very important part of using numerical libraries like numpy is **vectorising operations**: avoiding explicit loops over values in the arrays and instead using library functions to do manipulations. It is *massively* faster to have numpy add to arrays together than to iterate over the elements adding them together in Python.

### Good: vectorized

```python
a = np.zeros((5, 5))
b = np.ones((5, 5))
y = np.zeros((5, 5))

# YES: do things like this
# a and b are NumPy arrays
x = a + b
```

### Bad: explicit iteration

```python
# NO: don't do this
# this is very inefficient
for i in range(a.shape[0]):
    for j in range(a.shape[1]):
        y[i, j] = a[i, j] + b[i, j]
```

## No for loops (unless specified)
In this lab **do not** use explicit loops, like `while` or `for`, unless the question explicitly asks you to. **In future labs you may need to use the occasional `for` loop**, but try and avoid them where possible.

NumPy provides the **np.array** class which is a n-dimensional array of numbers **of the same type**. Arrays can be created in several ways: from a Python list (and nD arrays from nested lists), as a "blank" matrix of zeros or ones or random data, by copying existing arrays, loaded from disk or from certain special functions.

**You can always make a copy of an array using np.array() on an existing array** (e.g. `x = np.array(y)` makes a new **copy** of y). `np.array()` will also convert any iterable object (lists, tuples) into an array if it can. Note that a few operations will *change arrays in place*, and most will *return new copies*.

```python
x = np.array([1.0, 2.0, 3.0, 4.0])  # create an array from a list
print(x)  # print the array

# print the matrix, but nicer (note: the first string is a LaTeX expression)
print_matrix("x", x)

print(x.dtype)  # datatype
print(x.shape)  # shape of array
```

Output:
```
[1. 2. 3. 4.]
<IPython.core.display.Latex object>
float64
(4,)
```

# 1: Create some arrays
Create the following arrays, with the following specifications:
(**don't** use `np.array` to do this). If you don't know how to do this,
look at the worked example, or look at the cheatsheet or API reference.

Use `np.zeros`, `np.ones`, `np.full`, `np.arange` and `np.random.normal` to solve these questions.

* `x`: a 8 x 8 matrix of all zeros
* `y`: a 2 element vector, with all elements equal to `np.pi`
* `z`: a 1 x 2 x 5 element array of all ones.
* `q`: an 1D array with 10 elements, from 0-18 (inclusive), stepping by 2. 
* `r`: an 400 element array, random numbers normally distributed with mean 0, std. dev. 1.0 **(don't print this one out)**

Print out your arrays using `show_tensor_latex` to see if they look right.

Check that the tests pass. 

<!-- grade_id: cell-c190eec8e4831a17, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
x = np.zeros((8,8))
y = np.full((1,2), np.pi)[0]
z = np.ones((1,2,5))
q = np.arange(19, step=2)
r = np.random.normal(0,1.0,400)
```

<!-- grade_id: cell-7a8e23b3cc2d7bea, points: 1, grade: true, solution: false -->
```python
with marks(1):
    assert(check_hash(x,((8, 8), 0.0)))
```

Output:
```
<IPython.core.display.HTML object>
```

<!-- grade_id: cell-3eccab087d4c343c, points: 1, grade: true, solution: false -->
```python
with marks(1):
    assert(check_hash(y, ((2,),  21.991148575128552)))
```

Output:
```
<IPython.core.display.HTML object>
```

<!-- grade_id: cell-c3cd5694245364b2, points: 1, grade: true, solution: false -->
```python
with marks(1):
    assert(check_hash(z, ((1, 2, 5), 59.0)))
```

Output:
```
<IPython.core.display.HTML object>
```

<!-- grade_id: cell-f6fd4fc0f0e04fba, points: 1, grade: true, solution: false -->
```python
with marks(1):
    assert(check_hash(q,((10,), 701.744562646538)))
```

Output:
```
<IPython.core.display.HTML object>
```

<!-- grade_id: cell-8bc8bfb07bb942f3, points: 1, grade: true, solution: false -->
```python
with marks(1):
    assert(r.shape==(400,)  and np.std(r)>0.75 and np.std(r)<1.5 and np.all(np.diff(r.ravel())!=0.0))
```

Output:
```
<IPython.core.display.HTML object>
```

## 2: How to keep snails alive

<img src="imgs/snails.jpg" width="50%">*([Image](https://flickr.com/photos/chodhound/6083328289 "Snail") by [ChodHound](https://flickr.com/people/chodhound) license [CC BY-SA](https://creativecommons.org/licenses/by-sa/2.0/))*

Scientists at the Zoology Department, The University of Adelaide have studied the best conditions to keep snails alive. They have recorded a dataset of observations of snail mortality under controlled conditions. This data set is in the file `data/snails.txt`.

#### An excerpt from the data set description

    
>Groups of 20 snails were held for periods of 1, 2, 3 or 4 weeks in carefully
controlled conditions of temperature and relative humidity. There were two
species of snail, 0 and 1. At the end of the exposure time the snails
were tested to see if they had survived. 

>The data are unusual in that in most cases fatalities during the experiment
were fairly small. [lucky snails!]



### The task
The data is a 2D array, and has six columns, with these definitions:
            
            0              1                 2                3                 4          5
     species(binary) | exposure(weeks) | humidity(%) | temperature(deg. C) | n_deaths | n_snails
    
Each row represents one set of observations (i.e. one group of snails). You are to compute some basic properties of this data. Use NumPy operations to do the computations.

A. **Loading arrays** 
* Load this data as a NumPy array called `snails`. Note: use NumPy functions to do this! **Do not parse the file yourself** The file is space delimited.
* Print it out. Use this format to print out the results:
    
      print("snails\n", snails)

<!-- grade_id: cell-6d0228523852fbcb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
snails = np.loadtxt("./data/snails.txt", delimiter=" ")
```

<!-- grade_id: cell-d7f4d07872174036, points: 2, grade: true, solution: false -->
```python
with marks(2):
    assert(check_hash(snails, ((96, 6), 3082003.4024073719)))
```

Output:
```
<IPython.core.display.HTML object>
```

B. **Basic indexing** 
Compute the following results, storing the results in the variable specified and printing them out. Use the same printing format as A.


* `temp_first` the temperature in the first entry in the table. 
* `hum_diff` the *absolute* difference in humidity between the first and the last entry in the table.
* `weeks` the whole column of "weeks exposure".
* `row_fifth` the fifth row of observations. (remember 0 indexing)


<!-- grade_id: cell-17f2099d91b3038f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-3bd461f1be410f75, points: 1, grade: true, solution: false -->
```python

with marks(1):
    assert(check_hash(hum_diff, ((), 78.9)))
```

<!-- grade_id: cell-590385cd63ae790d, points: 1, grade: true, solution: false -->
```python
with marks(1):
    assert(check_hash(temp_first, ((), 50.0)))
```

<!-- grade_id: cell-19028943888ae090, points: 1, grade: true, solution: false -->
```python
with marks(1):
    assert(check_hash(weeks, ((96,), 13091.118033988751)))
    
```

<!-- grade_id: cell-57e22f24be8159a7, points: 1, grade: true, solution: false -->
```python
with marks(1):
    assert(check_hash(row_fifth, ((6,), 493.3630504844565)))
```

C. **Aggregate functions** 
Compute the following results, storing the results in the variable specified and printing them out:

* `total_deaths` total number of snails that died
* `total_still_alive` total number of snails that survived the whole study
* `mean_temp` mean temperature in the whole study
* `max_humidity` highest humidity in the study
* `average_death_rate` mean of the ratio of snail deaths to snails in the study
* `snail_weeks` the total amount of snail effort that went into this study (number of snails times number of weeks)

Each computation should be a single line of code

<!-- grade_id: cell-427021397b9a50b2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-12aba5ae7c698554, points: 2, grade: true, solution: false -->
```python
with marks(2):
    assert(check_hash(total_deaths, ((), 1375.0)))
    
```

<!-- grade_id: cell-0799d88a4e889521, points: 2, grade: true, solution: false -->
```python
with marks(2):
    assert(check_hash(total_still_alive, ((), 8225.0)))
```

<!-- grade_id: cell-e3030fa5597a6a3c, points: 2, grade: true, solution: false -->
```python
with marks(2):
    assert(check_hash(mean_temp, ((), 75.0)))
```

<!-- grade_id: cell-d7b5c4b686b8e185, points: 2, grade: true, solution: false -->
```python

with marks(2):
    assert(check_hash(max_humidity, ((), 379.0)))
```

<!-- grade_id: cell-759195cd947bae30, points: 3, grade: true, solution: false -->
```python
with marks(3):
    assert(check_hash(average_death_rate, ((), 0.71614583333333326)))
```

<!-- grade_id: cell-e75b606eb7b69e7a, points: 3, grade: true, solution: false -->
```python
with marks(3):
    assert(check_hash(snail_weeks, ((), 24000.0)))
```

D. **Boolean indexing**
Compute the following results, storing the results in the variable specified and printing them out:

* `species_0` and `species_1`: split the dataset into two arrays, one with the entries for species 0 and one with the entries for species 1.
    * Remember: Boolean arrays can be used *directly* as indices.

* `weakest_snail` the snail species (0 or 1) that had the highest average death rate

<!-- grade_id: cell-2fb2b916ca1d1254, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-33631e32555140e0, points: 5, grade: true, solution: false -->
```python
with marks(5):
    assert(check_hash(species_0, ((48, 6), 762041.58357993606)))
    assert(check_hash(species_1, ((48, 6), 791902.31693958596)))  
```

<!-- grade_id: cell-7405a52cf551031d, points: 3, grade: true, solution: false -->
```python
with marks(3):
    assert(check_hash(weakest_snail, ((), 5.0)))
```

E. **Arithmetic and ordering**
Compute the following results, storing the results in the variable specified and printing them out:

* `deg_f` each temperature in the study, but in degrees Fahrenheit. Use the knowledge that `0C = 32F, 100C = 212F`
* `mean_cols` the mean of all the columns, as a 1D vector
* `death_rate` the death rates, in sorted order, smallest first
* `exposure_death_order` the exposure durations (in weeks), but in sorted in the order of death rates, smallest death rate first.

* `best_temp`, `best_hum` the best temperature and humidity to keep a snail for four weeks without it dying. *Look only at the four week exposures, ignoring snails kept for less than this time.* 

<!-- grade_id: cell-965b0640474113ac, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-9aca90f5ed692ce6, points: 1, grade: true, solution: false -->
```python
with marks(1):
    assert(check_hash(deg_f, ((96,), 275523.34846922837)))
```

<!-- grade_id: cell-7d0f86f908ce2a87, points: 1, grade: true, solution: false -->
```python
with marks(1):
    assert(check_hash(mean_cols, ((6,), 522.92336963727359)))
```

<!-- grade_id: cell-e8a171753a823a96, points: 2, grade: true, solution: false -->
```python
with marks(2):
    assert(check_hash(death_rate, ((96,), 1100.9803982902313)))
```

<!-- grade_id: cell-554ca9fb59641ac9, points: 2, grade: true, solution: false -->
```python
with marks(2):
    assert(check_hash(exposure_death_order, ((96,), 13983.118033988751)))
```

<!-- grade_id: cell-22bc20db6ce8b95e, points: 3, grade: true, solution: false -->
```python
with marks(3):
    assert(check_hash(best_temp, ((), 50.0)))
```

<!-- grade_id: cell-36ed1c66c64c5946, points: 3, grade: true, solution: false -->
```python
with marks(3):
    assert(check_hash(best_hum, ((), 379.0)))      
```

## 3: Image operations
Images can be represented as numerical arrays. We will use images as an example to explore NumPy functionality.

* `img = ia.load_image_colour('filename.png')` will load an image as an array.
* `ia.show_image(img)` will show it in the notebook.
    

A)

We will:
* Load `data/parrots.png` as `img_array` 
* Print out its shape and dtype
* Show the image.

```python
img_array = ia.load_image_colour("data/parrots.png")
print(img_array.shape, img_array.dtype)
ia.show_image(img_array)
```

```python
copies = np.tile(img_array, (10, 1, 1, 1))
print(copies.shape)
ia.show_image(copies[0])
```

B) **Slicing arrays**
* Create an array `cropped` which has the pixels from [150,100] to [350,300]. Note that these positions are in `[row, col]` format, not `[x,y]`.
* Display the cropped array using `show_image()`. 
* Remember: the image is `WxHx3`. Think about how to slice the last dimension.
* Show the cropped image so you can see it.

<!-- grade_id: cell-e55532295e444bd0, points: 0, grade: false, solution: true -->
```python

# YOUR CODE HERE
```

<!-- grade_id: cell-1065728e6cf70cf4, points: 4, grade: true, solution: false -->
```python
ia.show_image(cropped)
with marks(4):
    assert(check_hash(cropped, ((200, 200, 3), 3409234926.1084023)))
```

C)  **Modifying arrays**

Create an array "censored" which is the same as `img_array`, but has a black bar across the following regions to protect the parrot's privacy:

    * [200,100] -> [260, 310]
    * [140, 400]-> [200, 650]

Setting array elements to zero will make them black.

**Do not modify the original `img_array`**

<!-- grade_id: cell-6615e422ecee6b7e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-fc66e323769d8a84, points: 4, grade: true, solution: false -->
```python
ia.show_image(censored)
with marks(4):
    assert(check_hash(censored, ((512, 768, 3), 250654064351.74332)))
    assert(check_hash(img_array, ((512, 768, 3), 269458072078.68924)))
```

D) **Stacking arrays**

Using *one* `for` loop, create a *list* of arrays which represent a panning animation which is a sequence of cropped versions of `img_array`. 

* The crops should start at [190,30] and be 200 pixels wide and 100 pixels high, and move 12 pixels right and 2 pixels up in each frame (that means the column index increases by 12, and the row index decreases by 2). 
* Create 30 frames. All of them will be the same size.
* Stack the array into a single array called `panning_array`.

Your code should be less than 15 lines (excluding any comments)!

(Optional extra: Work out how to use the [PIL](https://pillow.readthedocs.io/en/stable/) package to add a comedy caption to the animation).


<!-- grade_id: cell-3d8d7655e04e7662, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-7b663903137cc2cd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-3d8d7655e04e7660, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-b0bf82f051dbe3ef, points: 10, grade: true, solution: false -->
```python
ia.show_gif(panning_array, width="80%")

with marks(10):
    assert(check_hash(panning_array,((30, 100, 200, 3), 834973352307.7242)))
```

-----

# Submission instructions

## Before submission

* Make sure you fill in any place that says `YOUR CODE HERE` or `"YOUR ANSWER HERE"`.
* SAVE THE NOTEBOOK





<div class="alert alert-block alert-danger">
    
### Formatting the submission
* **WARNING**: If you do not submit the correct file, you will not get any marks.
* Submit this file **only** on Moodle. It will be named `<xxx>.ipynb`.

</div>


## Penalties (only for assessed labs)
<font color="red">
    
**Malformatted submissions**
</font>
These assignments are processed with an automatic tool; failure to follow instructions *precisely* will lead to you automatically losing two bands in grade regardless of whether the work is correct (not to mention a long delay in getting your work back). **If you submit a file without your work in it, it will be marked and you will get 0 marks.**

<font color="red">**Late submission**</font>
Be aware that there is a two band penalty for every *day* of late submission, starting the moment of the deadline.

<font color="red">
    
**Plagiarism**
</font> Plagiarism will be subject to the Plagiarism Policy. The penalties are severe.

```python

```